```python
# Initialize Otter
import otter
grader = otter.Notebook("Coding_Exercise_4.ipynb")
```

# Week 4 Coding Exercise

### Question 1
Write a function to calculate wage based on pay rate and hours worked:
- Unlike previous coding exercise, the function names and parameters have been written for you. For this exercise, you have to create the functions from scratch. Make sure the spelling and case are precise. Otherwise, the grader will not find your function and the tests will fail.
- Name the function **calculate_wage**
- The function has 2 parameters: **pay_rate** and **hours**
- Set the default value of **hours** to be 40
- First check that both the **hours** and **pay_rate** are positive. Otherwise, **return -1**. Remember that a function ends its execution when the return statement is executed. In other words, the function ends when negative values are passed in.
- If **hours** entered is <u>at or under 40</u>, pay with the **pay_rate**
- If **hours** entered is over 40, <u>the part exceeding 40</u> must be paid 1.5 times of the **pay_rate**
- The function <font color='red'>MUST RETURN the calculated wage</font>. You will get only half the credit if the wage is not returned even if the calculation is correct
- Make sure your code not only pass the public test cases but also has all the functionalities as instructed.

```python
...
### don't remove anything. Just add code to where noted
### Your code starts (you need to start with function definition)
def calculate_wage(pay_rate, hours=40):
    if pay_rate < 0 or hours < 0:
        return -1
    elif hours <= 40:
        return pay_rate * hours
    else:
        return pay_rate * 40 + (pay_rate * 1.5 * (hours - 40))
    
    ### Your code ends (don't forget to return your calculation. Don't print it out)
# You can call your function below to check your calculation but remove or comment the code before submission
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: passed
q1 - 2 message: passed
q1 - 3 message: passed
q1 - 4 message: passed
```

### Question 2
Write a program to open the text file "mbox_short.txt" (the file is provided with this coding exercise). If you download it multiple times, your computer will start to append (1), (2) to the file name. You need to make sure that your code read mbox_short.txt, not mbox_short(1).txt, etc. Otherwise, the test cases will fail. Open the file mbox_short.txt in a text editor like notepad to get familiar with the content before coding. In your code, you also need to open a file named 'output.txt' for the purpose of writing to it. An accumulator variable **total_lines** has been initiated for you. It is used to keep count of the number of lines in file mbox_short.txt. In a for loop, take each line in the file mbox_short.txt, increase the **total_lines** by 1. If it is a line with even count number, write it to the file output.txt. Outside the loop, close both files and RETURN **total_lines**. The return code has been written for you already. Just don't remove or change it.

```python
...
### don't remove anything. Just add code to where noted
def file_read_write():
    total_lines = 0
    ### Your code starts
    with open('mbox_short.txt', 'r') as source_file, open('output.txt', 'w') as dest_file:
        for line in source_file:
            even_line_bool = total_lines % 2 == 0
            if even_line_bool:
                dest_file.write(line)
            total_lines += 1
    ### Your code ends
    return total_lines
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# file_read_write()
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: correct
q2 - 2 message: correct
q2 - 3 message: correct
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Double check that the output.txt file is included in the zip file to be submitted.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True, files=['output.txt'])
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: passed
    q1 - 2 message: passed
    q1 - 3 message: passed
    q1 - 4 message: passed

    q2 results: All test cases passed!
    q2 - 1 message: correct
    q2 - 2 message: correct
    q2 - 3 message: correct
<IPython.core.display.HTML object>
```

Tests for q1 (points: 50)
```
>>> assert calculate_wage(20) == 800
```
failure_message: When pay_rate is 20 and the default hours is used, what should be the wage?
success_message: passed
```
>>> assert calculate_wage(20, -4) == -1
```
failure_message: If negative arguments are passed in, what should be returned
success_message: passed
```
>>> assert calculate_wage(-10) == -1
```
failure_message: If negative arguments are passed in, what should be returned
success_message: passed
```
>>> assert calculate_wage(10, 50) == 550
```
failure_message: If the work hours are higher than 40, you need to pay 1.5 times the pay rate for the extra hours
success_message: passed


Tests for q2 (points: 50)
```
>>> with open('mbox_short.txt', 'r') as f:
...     line = f.readline()
>>> assert len(line) > 0
```
failure_message: Is your mbox_short.txt file in the same folder? Is the file name correct? Have you closed the file in your code?
success_message: correct
```
>>> count = file_read_write()
>>> assert count > 0
```
failure_message: Make sure your function returns a positive integer
success_message: correct
```
>>> import os
>>> if os.path.exists('output.txt'):
...     os.remove('output.txt')
>>> count = file_read_write()
>>> with open('output.txt', 'r') as f:
...     lines = f.readlines()
>>> assert len(lines) > 0
```
failure_message: Is your output.txt file empty? Is the file name correct? Have you closed the file in your code?
success_message: correct
