```python
# Initialize Otter
import otter
grader = otter.Notebook("Coding_Exercise_6.ipynb")
```

# Week 6 Coding Exercise

### Question 1
Parse string to get email addresses in the file mbox_short.txt. Open the file mbox_short.txt in a text editor like notepad to get familiar with the text and email addresses before coding.
- The file mbox_short.txt contains email addresses. Your code should read the file, go through each line in the file to found the lines with email addresses. Add them to a list named **from_addresses**. The variable has been created as a empty list for you already.
    - Convert each line to <u>uppercase</u> before checking. The lines with email addresses now starts with "FROM:". An example line looks like this before being turned to uppercase: 
        - From: louis@media.berkeley.edu
    - When the loop is over, remember to close the file.
- In addition to **from_addresses**, three additional variables have been created for you at the top of the function to store data:  **domains** is a list to store email domains. **number_iupui** stores the number of addresses with 'IUPUI.EDU' as the domain. **number_umich** stores the number of addresses with 'UMICH.EDU' as the domain. 
- You should have noticed that the domain of an email is the part after the character @ in an email address. Use another for loop to iterate through the list **from_addresses**. Parse out the domain of each email and add it to the list **domains**.
- Now you can use methods (functions) of list to get the numbers of domains that are 'IUPUI.EDU' and those that are 'UMICH.EDU'. Save them to **number_iupui** and **number_umich** respectively.
- Return **from_addresses**, **domains**, **number_iupui** and **number_umich** in this order (this is already written out for you. Don't delete the code). Now you know the values will be returned as a tuple with 4 elements.
- Make sure your code pass the public test cases
- Test thoroughly for hiddent test cases

```python
...
### don't remove anything. Just add code to where noted
def analyze_email():
    from_addresses = []
    domains = []
    number_iupui = 0
    number_umich = 0
    ### Your code starts 
    with open('mbox_short.txt', 'r') as file:
        for line in file:
            line = line.strip().upper()
            if line.startswith("FROM:"):
                email = line[len("FROM: "):]
                from_addresses.append(email)
    
    for email in from_addresses:
        domain = email.split('@')[1]
        domains.append(domain)
        if domain == 'IUPUI.EDU':
            number_iupui += 1
        elif domain == 'UMICH.EDU':
            number_umich += 1
    
    ### Your code ends (the next line returns the values for the function. Don't remove it.)
    return from_addresses, domains, number_iupui, number_umich
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# print(analyze_email()[0])
# print(analyze_email()[1])
# print(analyze_email()[2])
# print(analyze_email()[3])
```

Output:
```
['STEPHEN.MARQUARD@UCT.AC.ZA', 'LOUIS@MEDIA.BERKELEY.EDU', 'ZQIAN@UMICH.EDU', 'RJLOWE@IUPUI.EDU', 'ZQIAN@UMICH.EDU', 'RJLOWE@IUPUI.EDU', 'CWEN@IUPUI.EDU', 'CWEN@IUPUI.EDU', 'GSILVER@UMICH.EDU', 'GSILVER@UMICH.EDU', 'ZQIAN@UMICH.EDU', 'GSILVER@UMICH.EDU', 'WAGNERMR@IUPUI.EDU', 'ZQIAN@UMICH.EDU', 'ANTRANIG@CARET.CAM.AC.UK', 'GOPAL.RAMASAMMYCOOK@GMAIL.COM', 'DAVID.HORWITZ@UCT.AC.ZA', 'DAVID.HORWITZ@UCT.AC.ZA', 'DAVID.HORWITZ@UCT.AC.ZA', 'DAVID.HORWITZ@UCT.AC.ZA', 'STEPHEN.MARQUARD@UCT.AC.ZA', 'LOUIS@MEDIA.BERKELEY.EDU', 'LOUIS@MEDIA.BERKELEY.EDU', 'RAY@MEDIA.BERKELEY.EDU', 'CWEN@IUPUI.EDU', 'CWEN@IUPUI.EDU', 'CWEN@IUPUI.EDU']
['UCT.AC.ZA', 'MEDIA.BERKELEY.EDU', 'UMICH.EDU', 'IUPUI.EDU', 'UMICH.EDU', 'IUPUI.EDU', 'IUPUI.EDU', 'IUPUI.EDU', 'UMICH.EDU', 'UMICH.EDU', 'UMICH.EDU', 'UMICH.EDU', 'IUPUI.EDU', 'UMICH.EDU', 'CARET.CAM.AC.UK', 'GMAIL.COM', 'UCT.AC.ZA', 'UCT.AC.ZA', 'UCT.AC.ZA', 'UCT.AC.ZA', 'UCT.AC.ZA', 'MEDIA.BERKELEY.EDU', 'MEDIA.BERKELEY.EDU', 'MEDIA.BERKELEY.EDU', 'IUPUI.EDU', '
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: passed - the returned values are in a tuple
q1 - 2 message: passed - the returned tuple has 4 elements
q1 - 3 message: passed - you found at least one email address
q1 - 4 message: passed - you added at least one item in the list of domains
q1 - 5 message: passed - the first from address found is correct and in upper case
```

### Question 2
Parse string to get weekdays in the file mbox_short.txt as described below. Open the file mbox_short.txt in a text editor like notepad to get familiar with the text and weekdays before coding:
- Your code should read the file, go through each line in the file to find the lines starting with "From:". Don't change the letter case in this question. Such lines have an email address following "From:". Some of them also have a datetime following the email address. Some of them don't have a datetime. Two such examples are provided below:
    - A line with datetime: From louis@media.berkeley.edu Fri Jan  4 18:10:48 2008
    - A line without datetime: From: louis@media.berkeley.edu
- check whether the line has weekday infromation. If so, the information should be recorded in a dictionary in a variable **count_dict**, which has been created for you as an empty dictionary at the begining of the function. Don't remove it. Use it instead.
- The goal is to store in **count_dict** the number of times any weekday appears in the lines starting with "From:". If a line does not have weekday, add 1 to the key "Unknown" in the dictionary. So the dictionary should look like the following if only Monday (4 times) and Tuesday (2 times) are in the lines and the rest don't have weekdays recorded (10 times). The solution will be different from the following but you have the idea now. You will see that the short weekday name appears consistently in the file. No additional conversion is needed.
    -{'Mon': 4, 'Tue': 2, 'Unknown': 10}
- When the loop is over, remember to close the file. Now **count_dict** should have all the info.
- Return **count_dict** and the weekday appeared the max number in the dictionary. A variable **weekday_max** was created for you also. It is possible that the max number is "Unknown". The two values will be returned as a tuple.
- Make sure your code pass the public test cases
- Test thoroughly for hiddent test cases also.

```python
...
### don't remove anything. Just add code to where noted
from datetime import datetime
def get_weekday():
    count_dict = {}
    weekday_max = ''
    ### Your code starts
    with open('mbox_short.txt', 'r') as file:
        for line in file:
            line = line.strip()
            if line.startswith("From ") or line.startswith("From:"):
                parts = line.split()
                # print(parts)
                if len(parts) >= 6:
                    weekday = parts[2]
                    if weekday in count_dict:
                        count_dict[weekday] += 1
                    else:
                        count_dict[weekday] = 1
                else:
                    if 'Unknown' in count_dict:
                        count_dict['Unknown'] += 1
                    else:
                        count_dict['Unknown'] = 1
    file.close()
    if count_dict:
        weekday_max = max(count_dict, key=count_dict.get)
    
    ### Your code ends
    return count_dict, weekday_max
# You can uncomment the lines below to check your return values but you should comment out the code before submission
print(get_weekday()[0])
print(get_weekday()[1])
```

Output:
```
{'Sat': 1, 'Unknown': 27, 'Fri': 20, 'Thu': 6}
Unknown
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: passed - your function returns a tuple with 2 elements
q2 - 2 message: passed -  the dictionary returned is not empty
q2 - 3 message: passed - The second returned value (the weekday with max appearance in the dictionary) is not an empty string
q2 - 4 message: passed - for lines without weekday info, you set the value to 'Unknown'
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Double check that the file mbox_short.txt is in the same folder as the notebook and the file is exactly as is. Test your code fully to pass hidden tests. This link leads to the video on how to test code: https://iu.mediaspace.kaltura.com/media/t/1_mmyb6sua

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 50)
```
>>> assert isinstance(analyze_email(), tuple)
```
failure_message: the function should return a tuple
success_message: passed - the returned values are in a tuple
```
>>> assert len(analyze_email()) == 4
```
failure_message: the function should return a tuple with 4 elements
success_message: passed - the returned tuple has 4 elements
```
>>> assert len(analyze_email()[0]) > 0
```
failure_message: Have you added any item to the list from_addresses?
success_message: passed - you found at least one email address
```
>>> assert len(analyze_email()[1]) > 0
```
failure_message: Have you added any item to the list domains?
success_message: passed - you added at least one item in the list of domains
```
>>> assert analyze_email()[0][0].find('STEPHEN.MARQUARD@UCT.AC.ZA') >= 0
```
failure_message: Have you located the correct lines to be added to the from_addresses? Have you converted the line to be upper case?
success_message: passed - the first from address found is correct and in upper case


Tests for q2 (points: 50)
```
>>> assert len(get_weekday()) == 2
```
failure_message: the function should return a tuple with 2 elements
success_message: passed - your function returns a tuple with 2 elements
```
>>> assert len(get_weekday()[0]) > 0
```
failure_message: The dictionary should not be empty
success_message: passed -  the dictionary returned is not empty
```
>>> assert get_weekday()[1] != ''
```
failure_message: The weekday with maximum appearance should not be an empty string
success_message: passed - The second returned value (the weekday with max appearance in the dictionary) is not an empty string
```
>>> assert 'Unknown' in get_weekday()[0]
```
failure_message: For lines without weekday info, did you count them to a specified key (instruction)? Check spelling and case also.
success_message: passed - for lines without weekday info, you set the value to 'Unknown'
