```python
# Initialize Otter
import otter
grader = otter.Notebook("Coding_Exercise_9.ipynb")
```

```python
import numpy as np
np.set_printoptions(suppress=True)
%precision 3
```

Output:
```
'%.3f'
```

# Week 9 Coding Exercise

### Question 1
Create a one-dimensional array using linspace(), start with 15 and end with 25. Generate 25 elements. Reshape it into a 5 X 5 array named **arr1**. Take the square of each element in **arr1** and assign the result back to **arr1**. Create another 5 X 5 array using random numbers by sampling from a normal distrubtion with the **mean** and **std** passed in to this function. The random **seed** should also be passed to the function. Use it to seed the random number generation. Name the second array **arr2**. Add **arr1** and **arr2** and <font color="red">RETURN</font> the result.

```python
...
### don't remove anything. Just add code to where noted
def get_array(mean, std, seed): 
    ### Your code starts 
    arr1 = np.linspace(15, 25, 25).reshape(5, 5)
    arr1 = np.square(arr1)
    
    np.random.seed(seed)
    arr2 = np.random.normal(mean, std, (5, 5))
    
    result = arr1 + arr2
    return result
   
    
    ### Your code ends (don't forget to return your calculation. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# get_array(10, 2, 1)
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: passed
q1 - 2 message: passed
```

### Question 2
Use the seed value passed to the function in the random number generation. Create a 6 X 5 array by getting random integers between 500 and 1000. Get the row means and column means and and <font color="red">RETURN</font> them in that order, meaing that row means should be returned first and then the column means. Note that the means shoudl be 2 1-dimensional arrays. Also remember that multiple return values will be in a tuple.

```python
...
### don't remove anything. Just add code to where noted
def get_means(seed): 
    ### Your code starts 
    np.random.seed(seed)
    arr = np.random.randint(500, 1000, (6, 5))
    
    row_means = np.mean(arr, axis=1)
    col_means = np.mean(arr, axis=0)
    
    return row_means, col_means
    
    ### Your code ends (don't forget to return your calculation. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# get_means(33)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: passed
q2 - 2 message: passed
```

### Question 3
Use the seed value passed to the function in the random number generation. Create a 6 X 5 array using random numbers with rand() function. Use Boolean indexing and slicing to get *columns* corresponding to labels 'c', or 'e' in the **label** array and <font color="red">RETURN</font> the array with the 2 columns. 

```python
...
### don't remove anything. Just add code to where noted
def get_columns(seed):
    label = np.array(['a', 'b', 'c', 'd', 'e'])
    ### Your code starts 
    np.random.seed(seed)
    arr = np.random.rand(6, 5)

    selected_columns = arr[:, (label == 'c') | (label == 'e')]
    
    return selected_columns
   
    
    ### Your code ends (don't forget to return your calculation. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# get_columns(1)
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
q3 - 1 message: passed
```

### Question 4
Use the **seed** value passed to the function in the random number generation. create an 6X6 array by using random numbers sampling from a normal distrubtion with the mean 100 and standard deviation 10. The function also has two other parameters: **upper** and **lower**. Assume that upper > lower (you don't have to check it) Replace the values in the array with **upper** if they are greater than it and **lower** if they are smaller than it. Other values should stay the same. <font color="red">RETURN</font> the updated array.

```python
...
### don't remove anything. Just add code to where noted
def cap_off(upper, lower, seed): 
    ### Your code starts 
    np.random.seed(seed)
    arr = np.random.normal(100, 10, (6, 6))
    
    arr = np.where(arr > upper, upper, arr)
    arr = np.where(arr < lower, lower, arr)
    
    return arr
    
    ### Your code ends (don't forget to return your calculation. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# cap_off(115, 85, 1)
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
q4 - 1 message: passed
q4 - 2 message: passed
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: passed
    q1 - 2 message: passed

    q2 results: All test cases passed!
    q2 - 1 message: passed
    q2 - 2 message: passed

    q3 results: All test cases passed!
    q3 - 1 message: passed

    q4 results: All test cases passed!
    q4 - 1 message: passed
    q4 - 2 message: passed
<IPython.core.display.HTML object>
```

Tests for q1 (points: 25)
```
>>> assert type(get_array(10, 2, 1)) == np.ndarray
```
failure_message: get_array() must return a numpy array
success_message: passed
```
>>> assert get_array(10, 2, 1).size == 25
```
failure_message: get_array() should have 25 values in the returned array.
success_message: passed


Tests for q2 (points: 25)
```
>>> assert type(get_means(1)) == type(())
```
failure_message: get_means() must return a tuple
success_message: passed
```
>>> assert get_means(1)[0].shape == (6,)
```
failure_message: there should be 6 row means
success_message: passed


Tests for q3 (points: 25)
```
>>> assert get_columns(1).shape == (6, 2)
```
failure_message: get_columns() should return a 6X2 array
success_message: passed


Tests for q4 (points: 25)
```
>>> assert cap_off(115, 85, 1).shape == (6, 6)
```
failure_message: cap_off() returns a 6X6 array
success_message: passed
```
>>> assert np.isclose(cap_off(115, 85, 1)[0, 1:5], np.array([93.882, 94.718, 89.27, 108.654]), 0.001).all()
```
failure_message: are you using the correct random function to generate the array?
success_message: passed
