```python
# Initialize Otter
import otter
grader = otter.Notebook("Coding_Exercise_3.ipynb")
```

# Week 3 Coding Exercise
<font color="red">Please don't insert or delete any cells on this notebook. Don't delete any existing text in any cells either </font>

### Question 1
Write a for loop in the following cell to see how many multiples of 7 are between two variables **lower_int** and **higher_int**:
- The acumulator variable **total_multiples_of_7** has been created for you already
- Use a for loop to check every integer between **lower_int** and **higher_int**. If it is a multiple of 7 (meaning that the remainder of dividing that number by 7 is 0. For example, 7, 14, 21, etc. are multiples of 7), add this integer to  **total_multiples_of_7**
    - Use % to check whether a number is a multiple of 7
    - Both **lower_int** and **higher_int** should be also be checked. In other words, they are included
    - If lower_int is 7 and higher_int is 35, the **total_multiples_of_7** should be 5 when the loop completes
- Assume **lower_int** and **higher_int** will be integers and **lower_int** <= **higher_int**. These two arguments will be set for each test so you don't need to set any values in your code. Just use them as variables.

```python
...
### don't remove anything. Just add code to where noted
def multiples_of_7(lower_int, higher_int):
    total_multiples_of_7 = 0
    ### Your code starts (write a for loop. Use modulo % to find if a number is multiple of 7)
    for i in range(lower_int, higher_int+1):
        if i % 7 == 0:
            total_multiples_of_7 += 1
    
    ### Your code ends
    return total_multiples_of_7
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# If you write print() in the code to check your work, you should remove them or comment them out before submission
# multiples_of_7(7, 35) # you should get 5
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: passed
q1 - 2 message: passed
q1 - 3 message: passed
```

### Question 2
We did an exercise in class to print out 4 lines of integers given an integer 4 like the following: \
1 \
1 2 \
1 2 3 \
1 2 3 4 \
Now, instead of print out the numbers, get the total of all the numbers that would be displayed. And program it to be flexible so if a variable **integer** is used to hold the integer, the total would change base on the value stored in **integer**. For example, 
- for number 1, the total is 1
- for number 2, the total is 4
- for number 3, the total is 10
- etc.

Tip: Nested for loop is most suitable to solve this problem. We worked in class to write the nested for loop. You can use the for loop directly. Just add all the numbers together.

Assume **integer** wil be an integer. This argument will be set for each test so you don't need to set any value in your code. Just use it as a variable.

```python
...
### don't remove anything. Just add code to where noted
def get_accumulated_total(integer):
    total = 0
    ### Your code starts
    for i in range(1, integer + 1):
        total += sum(range(1, i + 1))
    return total 
    
    ### Your code ends
    return total
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# If you write print() in the code to check your work, you should remove them or comment them out before submission
# get_accumulated_total(3) # the output should be 10
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: correct
q2 - 2 message: correct
q2 - 3 message: correct
```

### Question 3
Find all the numbers that are power of 2 (such as 4, 8, 16, etc.) using a while loop. The numbers found should be less than or equal to 5000. Concatenate all the power of 2 found in one string like the following in <ins>one line seperated by spaces</ins>: <pre> 2 4 8 16 32 64 128 256 512 1024 2048 4096 </pre> 
Tips:
- Use a while loop
- Two variables needed outside of the loop have been created for you already. **power_of_2** should increase by multipling by 2 with each loop iteration. **numbers_str** is used to hold all the powers of 2 with concatenation in each iteration.
- Remember string concatenation uses the + operator. All powers of 2 within 5000 should be displayed in a string on one line.

```python
...
### don't remove anything. Just add code to where noted
def get_power_of_2():
    power_of_2 = 2 
    numbers_str = ""
    ### Your code starts
    while power_of_2 <= 5000:
        numbers_str += str(power_of_2) + " "
        power_of_2 *= 2
    
    ### Your code ends
    return numbers_str
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# If you write print() in the code to check your work, you should remove them or comment them out before submission
# get_power_of_2() # should get '2 4 8 16 32 64 128 256 512 1024 2048 4096 '
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
q3 - 1 message: the total number of power of 2 found <=5000 is correct
q3 - 2 message: the smallest power of 2 found is correct
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: passed
    q1 - 2 message: passed
    q1 - 3 message: passed

    q2 results: All test cases passed!
    q2 - 1 message: correct
    q2 - 2 message: correct
    q2 - 3 message: correct

    q3 results: All test cases passed!
    q3 - 1 message: the total number of power of 2 found <=5000 is correct
    q3 - 2 message: the smallest power of 2 found is correct
<IPython.core.display.HTML object>
```

Tests for q1 (points: 30)
```
>>> assert multiples_of_7(0, 14) == 3
```
failure_message: when lower_int is 0 and higher_int is 14, how many multiples of 7 should you get?
success_message: passed
```
>>> assert multiples_of_7(1, 5) == 0
```
failure_message: when lower_int is 1 and higher_int is 5, how many multiples of 7 should you get?
success_message: passed
```
>>> assert multiples_of_7(1, 100) == 14
```
failure_message: when lower_int is 1 and higher_int is 100, how many multiples of 7 should you get?
success_message: passed


Tests for q2 (points: 30)
```
>>> assert get_accumulated_total(1) == 1
```
failure_message: incorrect result for integer 1
success_message: correct
```
>>> assert get_accumulated_total(2) == 4
```
failure_message: incorrect result for integer 2
success_message: correct
```
>>> assert get_accumulated_total(3) == 10
```
failure_message: incorrect result for integer 3
success_message: correct


Tests for q3 (points: 40)
```
>>> numbers = get_power_of_2().split()
>>> assert len(numbers) == 12
```
failure_message: the total number of power of 2 found <=5000 is incorrect
success_message: the total number of power of 2 found <=5000 is correct
```
>>> numbers = get_power_of_2().split()
>>> assert numbers[0] == '2'
```
failure_message: the smallest power of 2 found is incorrect
success_message: the smallest power of 2 found is correct
