```python
# Initialize Otter
import otter
grader = otter.Notebook("Coding_Exercise_7.ipynb")
```

# Week 7 Coding Exercise

### Question
Python variable names have the following rules:
1. Must start with a letter or underscore _
1. Must consist of letters, numbers, and underscores
1. Cannot contain spaces
1. Cannot use Python key (reserved) words
1. Max length is 79 characters
1. Case sensitive

Write a function to pull out valid variable names in a string passed in as an argument. Assume that the candidate variable names are seperated by a comma and a space. The function is named **get_valid_var_names** which takes one argument named **candidate_name_str**. This function header is already written for you. The function should put all valid variable names in the argument into a list named **valid_name_list** and have the list returned.

Part of the code has been provided for you already. You can follow the following steps to implement the function:
- split **candidate_name_str** into names using split() in re module. Candidate names are seperated by a comma and a space
- go through the list of reserved_words to remove all the reserved words in the names. No regular expression is needed for this step but be aware that remove() only removes one item in a list. If there are duplicates, only the first one is removed. Make sure all occurances of all keywords are removed even if some appear multiple times
- Create a regex that would check all the remaining names to pick out valid ones. Add the valid variable name to **valid_name_list**. You must use re module to get this done. Use a for loop to make processing easier.
- Return **valid_name_list**. This is already written for you. Just don't remove it.
- Make sure you test your code thoroughly.

```python
...
### don't remove anything. Just add code to where noted
import re
reserved_words = ['False', 'None', 'True', 'and', 'as', 'assert', 'async', 'await', 'break', 'class', 'continue', 'def', 'del', 'elif', 'else', 'except', 'finally', 'for', 'from', 'global', 'if', 'import', 'in', 'is', 'lambda', 'nonlocal', 'not', 'or', 'pass', 'raise', 'return', 'try', 'while', 'with', 'yield']
def get_valid_var_names(candidate_name_str):
    valid_name_list = []
    ### Your code starts (you need to start with function definition)
    candidate_names = candidate_name_str.split(", ")
    candidate_names = [name for name in candidate_names if name not in reserved_words]
    var_name_pattern = re.compile(r'^[A-Za-z_][A-Za-z0-9_]{0,78}$')
    for name in candidate_names:
        if var_name_pattern.match(name):
            valid_name_list.append(name)
    
    ### Your code ends (don't forget to return your calculation. Don't print it out)
    return valid_name_list
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# Feel free to test with additional test cases to cover all the bases. The hidden test cases could reveal problems
print(get_valid_var_names('')) # should get []
print(get_valid_var_names('1')) # should get []
print(get_valid_var_names('abc, 0_er, if, if')) # should get ['abc']
print(get_valid_var_names('abc, email, e 1')) # should get ['abc', 'email']
```

Output:
```
[]
[]
['abc']
['abc', 'email']
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: passed
q1 - 2 message: passed
q1 - 3 message: passed
q1 - 4 message: passed
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: passed
    q1 - 2 message: passed
    q1 - 3 message: passed
    q1 - 4 message: passed
<IPython.core.display.HTML object>
```

Tests for q1 (points: 100)
```
>>> assert len(get_valid_var_names('1')) == 0
```
failure_message: Python variable names cannot start with numbers
success_message: passed
```
>>> assert len(get_valid_var_names('a, y&n')) == 1
```
failure_message: Python variable names cannot contain chars other than digits, letters and underscore
success_message: passed
```
>>> assert len(get_valid_var_names('a, if, and, if, b1')) == 2
```
failure_message: Python variable names should not be reserved words (keywords)
success_message: passed
```
>>> assert len(get_valid_var_names('')) == 0
```
failure_message: Python variable names are case sensitive. A reserved word with a different case is a valid name
success_message: passed
