```python
# Initialize Otter
import otter
grader = otter.Notebook("Coding_Exercise_7.ipynb")
```

# Week 7 Coding Exercise

## Question - Get Valid Variable Names of Python
Python variable names have the following rules:
1. Must start with a letter or underscore _ 
1. Must consist of letters, numbers, and underscores
1. Cannot contain spaces
1. Cannot use Python key (reserved) words
1. Max length is 79 characters
1. Case sensitive

The reserved words are: 
<pre>
False, None, True, and, as, assert, async, await, break, class, continue, def, del, elif, else, except, finally, for, from, global, if, import, in, is, lambda, nonlocal, not, or, pass, raise, return, try, while, with, yield
</pre>

Write a function to pull out valid variable names in a string passed in. Assume that the possible variable names are seperated by a comma and a space. The function should be named **get_valid_var()** and it takes one argument named **text**. It should **return** a list of valid variable names in the argument passed in. The code cell has done the following for you already:
1. Imported the regular expression module re for you. 
1. Created a tuple of Python reserved words you can use to check possible variable names

Note in this coding exercise, you have to write the function header and return statement. In previous coding exercises, these are usually written for you. Here are the steps to implement the solution:
1. step 1, split the text into words using split() in re module assuming the text will have var names seperated by a comma and a space (see test_text in next cell for example)
1. step 2, go through the list of reserved_words to remove all the reserved words in the names. No regular expression is needed at this step but be aware that remove() only removes one item in a list. If there are duplicates, only the first one is removed. Make sure all occurances of all keywords are removed even if some appear multiple times
1. Step 3, Create a regex that would check all the remaining variable names to pick out valid ones. You can create a new empty list to store valid variable names. You must use re module to get this done. Use a for loop to go thorough each remaining names.
1. Step 4, return the list of valid names. DON'T PRINT the result. If you print things out while coding, remove all print() before final testing and submission

#### Test your function when it is done.
If you use the string ' var 1, apple, eggs, await, apple2, _cheese, 2apple, $dollar, var.22, APPLE, ui8655, r_tyM0, if, if' to test the function, you should have the following output returned
<pre>
['apple', 'eggs', 'apple2', '_cheese', 'APPLE', 'ui8655', 'r_tyM0']
</pre>

```python
...
### don't remove anything. Just add code to where noted
# Import re module
import re
reserved_words = ('False', 'None', 'True', 'and', 'as', 'assert', 'async', 'await', 'break', 'class', 'continue', 'def', 'del', 'elif', 'else', 'except', 'finally', 'for', 'from', 'global', 'if', 'import', 'in', 'is', 'lambda', 'nonlocal', 'not', 'or', 'pass', 'raise', 'return', 'try', 'while', 'with', 'yield')
### Your code starts 
def get_valid_var(text):
    candidate_names = re.split(', ', text)
    candidate_names = [name for name in candidate_names if name not in reserved_words]
    valid_var_regex = re.compile(r'^[A-Za-z_][A-Za-z0-9_]*$')
    valid_name_list = []
    for name in candidate_names:
        if valid_var_regex.search(name) and len(name) <= 79:
            valid_name_list.append(name)
    return valid_name_list
    
### Your code ends (Don't forget to return your result)
# You can uncomment the lines below to check your return values but you should comment out the code before submission
test_text = 'var 1, apple, eggs, await, apple2, _cheese, 2apple, $dollar, var.22, APPLE, ui8655, r_tyM0, if, if'
returned_names = get_valid_var(test_text)
print(returned_names)  
```

Output:
```
['apple', 'eggs', 'apple2', '_cheese', 'APPLE', 'ui8655', 'r_tyM0']
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: passed
q1 - 2 message: passed
q1 - 3 message: passed
q1 - 4 message: passed
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Test your code fully to pass hidden tests. This link leads to the video on how to test code: https://iu.mediaspace.kaltura.com/media/t/1_mmyb6sua

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: passed - You have created the function get_valid_var() that takes one argument
    q1 - 2 message: passed - Your function get_valid_var() returns a list. The return value is correct
    q1 - 3 message: passed - Your function detects invalid names starting with numbers
    q1 - 4 message: passed - Your function detects invalid names containing special characters
    q1 - 5 message: passed - Your function detects invalid names as Python reserved words
<IPython.core.display.HTML object>
```

Tests for q1 (points: 100)
```
>>> temp = get_valid_var('')
>>> assert 1 == 1
```
failure_message: Have you created the function get_valid_var() that takes one argument? Check spelling and case
success_message: passed - You have created the function get_valid_var() that takes one argument
```
>>> assert isinstance(get_valid_var(''), list)
```
failure_message: Does the function get_valid_var() returns a list?
success_message: passed - Your function get_valid_var() returns a list. The return value is correct
```
>>> assert len(get_valid_var('9th')) == 0
```
failure_message: Does the function get_valid_var() detect invalid names starting with numbers?
success_message: passed - Your function detects invalid names starting with numbers
```
>>> assert len(get_valid_var('_eb&')) == 0
```
failure_message: Does the function get_valid_var() detect invalid names containing special characters?
success_message: passed - Your function detects invalid names containing special characters
```
>>> assert len(get_valid_var('if, else, def, for')) == 0
```
failure_message: Does the function get_valid_var() detect invalid names as Python reserved words?
success_message: passed - Your function detects invalid names as Python reserved words
