```python
# Initialize Otter
import otter
grader = otter.Notebook("Coding_Exercise_10.ipynb")
```

```python
# This cell load all the needed libraries for you. Don't add or remove anything. You don't need other packages either
from IPython.display import display, HTML
display(HTML("<style>.container { width:100% !important; }</style>"))
import numpy as np 
import pandas as pd 
# prevent numbers to be displayed in scientific notation
np.set_printoptions(precision=3, suppress=True)
%precision 3
```

Output:
```
<IPython.core.display.HTML object>
'%.3f'
```

# Week 10 Coding Exercise
The following cell load the data in the csv file to the dataframe 'df' for you. It also displays some basic information. Use it to understand the data. <font color='red'>DON'T DELETE THE CELL</font>

```python
# Read the file to a dataframe
df = pd.read_csv("netflix_show_ranking.csv")
# Display the top 5 rows to take a peek of the data
display(df.head())
# Use info() to get information such as names of columns, number of values not missing, and data types of columns
df.info()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[1], line 2
      1 # Read the file to a dataframe
----> 2 df = pd.read_csv("netflix_show_ranking.csv")
      3 # Display the top 5 rows to take a peek of the data
      4 display(df.head())

NameError: name 'pd' is not defined
```

### Question 1
The column "Days In Top 10" has the numbers of days a show is listed in top 10 per a particular date in column 'As of'. Find the largest number in the column "Days In Top 10". Return both the name of the show and the Days of the show in Top 10 (in this order). <font color='red'>Use **df** loaded in the cell at the begining of this notebook. Don't read the csv file and load the data again.</font>

```python
...
### don't remove anything. Just add code to where noted
def get_show_in_top10_longest():
    ### Your code starts 
    max_days_index = df['Days In Top 10'].idxmax()
    bestShow = df.loc[max_days_index, ['Title', 'Days In Top 10']]
    return tuple(bestShow)
    ### Your code ends (don't forget to return your result. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# get_show_in_top10_longest()
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: passed - the type of the return value is correct
q1 - 2 message: passed - 2 values are returned in a tuple
q1 - 3 message: passed - The first value in the returned tuple is a string
q1 - 4 message: passed - The second value in the returned tuple is an integer
```

### Question 2
Get the following statistics of the column 'Viewership Score' and return all of them in a tuple **in the order as listed below**
- mean
- standard deviation
- median
- highest
- lowest
- skew

<font color='red'>Use **df** loaded in the cell at the begining of this notebook. Don't read the csv file and load the data again.</font>

```python
...
### don't remove anything. Just add code to where noted
def get_stats_viewing_score():
    ### Your code starts 
    mean = df['Viewership Score'].mean()
    standard_deviation = df['Viewership Score'].std()
    median = df['Viewership Score'].median()
    highest = df['Viewership Score'].max()
    lowest = df['Viewership Score'].min()
    skew = df['Viewership Score'].skew()
    return (mean, standard_deviation, median, highest, lowest, skew)
   
    
    ### Your code ends (don't forget to return your result. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
get_stats_viewing_score()
```

Output:
```
(160.730, 265.884, 58.000, 1437, 1, 2.899)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: passed - the type of the return value is correct
q2 - 2 message: passed - 6 values are returned in a tuple
q2 - 3 message: passed - The first value in the returned tuple is an integer
q2 - 4 message: passed - The last value in the returned tuple is a float
```

### Question 3
Find the number of shows/titles that show up in top 10 ranking **only once** in 2021 (the time period of the dataset). Return this number. Hint: use value_counts() and work on the returned value of this method

<font color='red'>Use **df** loaded in the cell at the begining of this notebook. Don't read the csv file and load the data again.</font>

```python
...
### don't remove anything. Just add code to where noted
def get_num_shows_top10_once():
    ### Your code starts 
    title_counts = df['Title'].value_counts()
    title_once = title_counts[title_counts == 1].count()
    return int(title_once)
    
    ### Your code ends (don't forget to return your result. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
get_num_shows_top10_once()
```

Output:
```
17
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
q3 - 1 message: passed - the type of the return value is correct
q3 - 2 message: passed - the return value is in the correct range
```

### Question 4
Get the title of the show that is ranked number 1 in the top 10 list the highest number of times. Note that the title instead of the number should be returned. Again, value_counts() should be helpful

<font color='red'>Use **df** loaded in the cell at the begining of this notebook. Don't read the csv file and load the data again.</font>

```python
...
### don't remove anything. Just add code to where noted
def get_show_most_top1():
    ### Your code starts 
    top1titles = df[df['Rank'] == 1]
    top1titles = top1titles['Title'].value_counts()
    top1titles = top1titles.idxmax()
    return top1titles
   
    
    ### Your code ends (don't forget to return your result. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# get_show_most_top1()
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
q4 - 1 message: passed - the type of the return value is correct
q4 - 2 message: passed - the number of words in the returned string is correct
```

### Question 5
Find all the rows that are "Netflix Exclusive" **movies** that have been ranked number 1. Take the result dataframe and only return 3 columns 'Title', 'Days In Top 10', 'Viewership Score' and sorted by 'Viewership Score' in descending order

<font color='red'>Use **df** loaded in the cell at the begining of this notebook. Don't read the csv file and load the data again.</font>

```python

```

```python
...
### don't remove anything. Just add code to where noted
def get_top1_movies_sorted():
    ### Your code starts 
    filtered_df = df[(df['Netflix Exclusive'] == 'Yes') & 
                     (df['Type'] == 'Movie') & 
                     (df['Rank'] == 1)]
    result_df = filtered_df[['Title', 'Days In Top 10', 'Viewership Score']]
    result_df = result_df.sort_values(by='Viewership Score', ascending=False)   
    return result_df
    
    ### Your code ends (don't forget to return your result. Don't print it out) 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# get_top1_movies_sorted()
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
q5 - 1 message: passed - the type of the return value is correct
q5 - 2 message: passed - the number of rows of the returned dataframe is correct
q5 - 3 message: passed - the number of columns of the returned dataframe is correct
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: passed - the type of the return value is correct
    q1 - 2 message: passed - 2 values are returned in a tuple
    q1 - 3 message: passed - The first value in the returned tuple is a string
    q1 - 4 message: passed - The second value in the returned tuple is an integer

    q2 results: All test cases passed!
    q2 - 1 message: passed - the type of the return value is correct
    q2 - 2 message: passed - 6 values are returned in a tuple
    q2 - 3 message: passed - The first value in the returned tuple is an integer
    q2 - 4 message: passed - The last value in the returned tuple is a float

    q3 results: All test cases passed!
    q3 - 1 message: passed - the type of the return value is correct
    q3 - 2 message: passed - the return value is in the correct range

    q4 results: All
```

Tests for q1 (points: 20)
```
>>> assert type(get_show_in_top10_longest()) == type(())
```
failure_message: get_show_in_top10_longest() must return a tuple
success_message: passed - the type of the return value is correct
```
>>> assert len(get_show_in_top10_longest()) == 2
```
failure_message: The returned tuple should contain 2 values
success_message: passed - 2 values are returned in a tuple
```
>>> assert isinstance(get_show_in_top10_longest()[0], str)
```
failure_message: The first value in the returned tuple should be a string
success_message: passed - The first value in the returned tuple is a string
```
>>> assert isinstance(get_show_in_top10_longest()[1], np.int64)
```
failure_message: The second value in the returned tuple should be an integer
success_message: passed - The second value in the returned tuple is an integer


Tests for q2 (points: 20)
```
>>> assert type(get_stats_viewing_score()) == type(())
```
failure_message: get_stats_viewing_score() must return a tuple
success_message: passed - the type of the return value is correct
```
>>> assert len(get_stats_viewing_score()) == 6
```
failure_message: The returned tuple should contain 6 values
success_message: passed - 6 values are returned in a tuple
```
>>> assert isinstance(get_stats_viewing_score()[0], np.float64)
```
failure_message: The first value in the returned tuple should be an integer
success_message: passed - The first value in the returned tuple is an integer
```
>>> assert isinstance(get_stats_viewing_score()[-1], np.float64)
```
failure_message: The last value in the returned tuple should be a float
success_message: passed - The last value in the returned tuple is a float


Tests for q3 (points: 20)
```
>>> assert isinstance(get_num_shows_top10_once(), int)
```
failure_message: get_num_shows_top10_once() must return a integer
success_message: passed - the type of the return value is correct
```
>>> assert get_num_shows_top10_once() < 50
```
failure_message: get_num_shows_top10_once() must return a integer < 50
success_message: passed - the return value is in the correct range


Tests for q4 (points: 20)
```
>>> assert isinstance(get_show_most_top1(), str)
```
failure_message: get_num_shows_top10_once() must return a string
success_message: passed - the type of the return value is correct
```
>>> assert len(get_show_most_top1().split()) == 2
```
failure_message: get_num_shows_top10_once() must return a string with 2 words
success_message: passed - the number of words in the returned string is correct


Tests for q5 (points: 20)
```
>>> assert isinstance(get_top1_movies_sorted(), pd.DataFrame)
```
failure_message: get_top1_movies_sorted() must return a Pandas dataframe
success_message: passed - the type of the return value is correct
```
>>> assert get_top1_movies_sorted().shape[0] == 113
```
failure_message: get_top1_movies_sorted() must return a Pandas dataframe with 113 rows
success_message: passed - the number of rows of the returned dataframe is correct
```
>>> assert get_top1_movies_sorted().shape[1] == 3
```
failure_message: get_top1_movies_sorted() must return a Pandas dataframe with 3 columns
success_message: passed - the number of columns of the returned dataframe is correct
