```python
# Initialize Otter
import otter
grader = otter.Notebook("Coding_Exercise_5.ipynb")
```

# Week 5 Coding Exercise

### Question 1
Parse string to get email senders in the file mbox_short.txt. Open the file mbox_short.txt in a text editor like notepad to get familiar with the text and email senders before coding.
- The file mbox_short.txt contains email senders. Your code should read the file, go through each line to pick the email senders. Add the senders in a string and keep count of how many senders are in the file. Some example lines with possible senders are listed below:
    - Received: from prod.collab.uhi.ac.uk ([194.35.219.182])
    - Received: from paploo.uhi.ac.uk (localhost [127.0.0.1])
- Both lines starts with "Received: from". The first line does not contain "localhost" so it will be considered as a sender. 
- The second line contains "localhost" so it will not be considered as a sender.
- This is how you can write the code:
    - Open the file, use a for loop to go through each line.
    - Turn the line to lower case so your search will be case insensitive
    - For the line that start with "received" and doesn't contain "localhost", get the part of the line after "received: from" as the sender. (Note: use lower case since you converted the line to be lower case) For the example above, the sender is:
        - prod.collab.uhi.ac.uk ([194.35.219.182])
    - Two variables have been created for you to store the information:  **total_senders** as the count of senders in the file. **senders_str** is to concatenate all the senders found seperated by a line break "\n"
    - If the line has a sender, concatenate the sender to **senders_str**. Concatenate a '\n' at the end as well
    - Increment **total_senders** by 1 
    - When the loop ends, close the file and return the two variables (this is already written out for you. Don't delete it)
- Make sure your code pass the public test cases

```python
...
### don't remove anything. Just add code to where noted
def get_senders():
    total_senders = 0
    senders_str = ''
    ### Your code starts (you need to start with function definition)
    file = open('mbox_short.txt')
    for line in file:
        line = line.lower().strip()
        if line.startswith("received: from") and 'localhost' not in line:
            total_senders += 1
            senders_str += line[len("received: from"):] + '\n'
    file.close()
    ### Your code ends (don't forget to return your calculation. Don't print it out)
    return total_senders, senders_str
print(get_senders()[1])
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# print(get_senders()[0])
# print(get_senders()[1])
```

Output:
```
murder (mail.umich.edu [141.211.14.90])
 murder ([unix socket])
 holes.mr.itd.umich.edu (holes.mr.itd.umich.edu [141.211.14.79])
 paploo.uhi.ac.uk (app1.prod.collab.uhi.ac.uk [194.35.219.184])
 prod.collab.uhi.ac.uk ([194.35.219.182])
 nakamura.uits.iupui.edu (nakamura.uits.iupui.edu [134.68.220.122])
 murder (mail.umich.edu [141.211.14.97])
 murder ([unix socket])
 icestorm.mr.itd.umich.edu (icestorm.mr.itd.umich.edu [141.211.93.149])
 paploo.uhi.ac.uk (app1.prod.collab.uhi.ac.uk [194.35.219.184])
 prod.collab.uhi.ac.uk ([194.35.219.182])
 nakamura.uits.iupui.edu (nakamura.uits.iupui.edu [134.68.220.122])
 murder (mail.umich.edu [141.211.14.25])
 murder ([unix socket])
 ghostbusters.mr.itd.umich.edu (ghostbusters.mr.itd.umich.edu [141.211.93.144])
 paploo.uhi.ac.uk (app1.prod.collab.uhi.ac.uk [194.35.219.184])
 prod.collab.uhi.ac.uk ([194.35.219.182])
 nakamura.uits.iupui.edu (nakamura.uits.iupui.edu [134.68.220.122])
 murder (mail.umich.edu [141.211.14.25])
 murder ([unix socket])
 d
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: passed
q1 - 2 message: passed
q1 - 3 message: passed
```

### Question 2
Parse string to get dates in the file mbox_short.txt. Open the file mbox_short.txt in a text editor like notepad to get familiar with the text and datetime before coding.
- The file mbox_short.txt contains dates. Your code should read the file, go through each line to pick the dates including the times. Add the dates and times in a string and keep count of how many dates are in the file. Two example lines with dates are provided below:
    - Date: Fri, 4 Jan 2008 16:09:02 -0500
    - Date: 2008-01-04 15:01:37 -0500 (Fri, 04 Jan 2008)
- Two variables have been created for you to store the information:  **total_datetime** as the count of dates in the file. **datetime_str** is to concatenate all the dates found seperated by a line break "\n"
- Here is how can can implement the code:
    - Open the file, use a for loop to go through each line.
    - Turn the line to lower case so your search will be case insensitive
    - For the line that starts with "date:", get the part of the line after "date: ". Use strptime() and a proper format string to convert the string to a datetime object (Note: use lower case since you converted the line to be lower case). There are 2 types of datetime formats in the file as can be seen in the two examples above. Different format strings are need for the two types.
        1. Fri, 4 Jan 2008 16:09:02
        1. 2008-01-04 15:01:37
    - Use strftime to get a string representation of the date in the format of yyyy-mm-dd hours:minutes:seconds. Then concatenate the string to **datetime_str**. Concatenate a '\n' at the end as well. The final format of the datetime should all be:
    - 2008-01-04 16:09:02
    - Increment **total_datetime** by 1 
    - When the loop ends, close the file and return the two variables (this is already written out for you. Don't delete it)
- Make sure your code pass the public test cases

```python
...
### don't remove anything. Just add code to where noted
from datetime import datetime
def get_datetime():
    total_datetime = 0
    datetime_str = ''
    ### Your code starts
    with open('mbox_short.txt') as file:
        for line in file:
            line = line.strip()
            if line.lower().startswith("date: "):
                date_str = line[6:]
                if '(' in date_str:
                    date_str = date_str.split('(', 1)[0].strip()
                if date_str.endswith('-0500') or date_str.endswith('+0000'):
                    date_str = date_str[:-6].strip()
                try:
                    date_obj = datetime.strptime(date_str, "%a, %d %b %Y %H:%M:%S")
                except ValueError:
                    try:
                        date_obj = datetime.strptime(date_str, "%Y-%m-%d %H:%M:%S")
                    except ValueError:
                        continue
                formatted_date = date_obj.strftime("%Y-%m-%d %H:%M:%S")
                datetime_str += formatted_date + '\n'
                total_datetime += 1
    ### Your code ends
    return total_datetime, datetime_str 
# You can uncomment the lines below to check your return values but you should comment out the code before submission
# print(get_datetime()[0])
# print(get_datetime()[1])
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: correct
q2 - 2 message: correct
q2 - 3 message: correct
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Double check that the file mbox_short.txt is in the same folder as the notebook and the file is exactly as is

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: passed
    q1 - 2 message: passed
    q1 - 3 message: passed

    q2 results: All test cases passed!
    q2 - 1 message: correct
    q2 - 2 message: correct
    q2 - 3 message: correct
<IPython.core.display.HTML object>
```

Tests for q1 (points: 50)
```
>>> with open('mbox_short.txt', 'r') as f:
...     line = f.readline()
>>> assert len(line) > 0
```
failure_message: Is mbox_short.txt file in the same folder? Is the file name correct? Have you closed the file in your code?
success_message: passed
```
>>> assert get_senders()[0] > 0
```
failure_message: The total number of senders should be more than 0
success_message: passed
```
>>> assert get_senders()[1] != ''
```
failure_message: The senders_str should not be an empty string
success_message: passed


Tests for q2 (points: 50)
```
>>> with open('mbox_short.txt', 'r') as f:
...     line = f.readline()
>>> assert len(line) > 0
```
failure_message: Is your mbox_short.txt file in the same folder? Is the file name correct? Have you closed the file in your code?
success_message: correct
```
>>> assert get_datetime()[0] > 0
```
failure_message: The total number of dates should be more than 0
success_message: correct
```
>>> assert get_datetime()[1] != ''
```
failure_message: The datetime_str should not be an empty string
success_message: correct
